---
title       : Termin 3 (homework) - Noch nicht verfügbar
description : Kontrollstrukturen und Funktionen in R
--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:a89b75f3c5
## test
test

*** =instructions
- 1
- 2

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
msg_bad <- "Leider falsch!"
msg_success <- "Richtig!"
test_mc(correct = 1, feedback_msgs = c(msg_bad, msg_success, msg_bad))

```