---
title       : Einheit 5 (homework)
description : Noch nicht verfügbar
--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:e76843b333
## <<<New Exercise>>>

xx
*** =instructions
- test1
- test2
*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
test_mc(correct = 1, feedback_msgs = c(msg_bad, msg_success, msg_bad))
```
